# Project Euler: Problem 8

https://projecteuler.net/problem=8

The four adjacent digits in the 1000-digit number that have the greatest product are 9 × 9 × 8 × 9 = 5832.

(see the number below)

Find the thirteen adjacent digits in the 1000-digit number that have the greatest product. What is the value of this product?

Use NumPy for this computation

```python
import numpy as np
```

```python
d1000 = 7316717653133062491922511967442657474235534919493496983520312774506326239578318016984801869478851843858615607891129494954595017379583319528532088055111254069874715852386305071569329096329522744304355766896648950445244523161731856403098711121722383113622298934233803081353362766142828064444866452387493035890729629049156044077239071381051585930796086670172427121883998797908792274921901699720888093776657273330010533678812202354218097512545405947522435258490771167055601360483958644670632441572215539753697817977846174064955149290862569321978468622482839722413756570560574902614079729686524145351004748216637048440319989000889524345065854122758866688116427171479924442928230863465674813919123162824586178664583591245665294765456828489128831426076900422421902267105562632111110937054421750694165896040807198403850962455444362981230987879927244284909188845801561660979191338754992005240636899125607176060588611646710940507754100225698315520005593572972571636269561882670428252483600823257530420752963450
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
e1000 = str(d1000)
c = 0
n = -1
x = 1
product = 0

a = np.ones((13),int) #See I used numpy
while c<1000: #The loop goes through the entire sequence
    c = c + 1
    n = n + 1
    #Each value of the array is shifted over 1, and the new value is placed at the 13th spot
    a[0] = a[1]
    a[1] = a[2]
    a[2] = a[3]
    a[3] = a[4]
    a[4] = a[5]
    a[5] = a[6]
    a[6] = a[7]
    a[7] = a[8]
    a[8] = a[9]
    a[9] = a[10]
    a[10] = a[11]
    a[11] = a[12]
    a[12] = int(e1000[n])
    x = np.cumprod(a)  #The product of the array
    x = x[12]
    if product < x:  #Replaces the product if it is larger than the largest so far.
        product = x
print (product)
print (x)
#raise NotImplementedError()
```

Output:
```
23514624000
0
```

<!-- grade_id: projecteuler8, points: 10, grade: true, solution: false -->
```python
assert True # leave this for grading
```