# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    z = []
    w = []
    n = 0
    x = str(x)
    y = str(y)
    thing = True
    while n < len(x):
        z.append(x[n])
        n = n+1
    n = 0
    while n < len(y):
        w.append(y[n])
        n = n+1
    z = sorted(z)
    w = sorted(w)

    return z == w
#same_digits('132', '321')
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
x = 0
while x < 10000000:
    x = x + 1
    if same_digits(x, 2*x) and same_digits(x, 3*x) and same_digits(x, 4*x) and same_digits(x, 5*x) and same_digits(x, 6*x):
        z = x
        break
print(z)
```

Output:
```
142857
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```